# Lab 01：数据科学实验导引
欢迎来到数据科学课程第一次实验。本实验帮助你熟悉 Jupyter Notebook、复习 Python 与 `NumPy`，并初步使用 `matplotlib` 进行可视化。

请按顺序完成所有题目并通过测试单元。建议完成后执行一次“Restart & Run All”检查结果是否可复现。

### 协作说明

你可以与同学讨论思路，但请**独立完成并书写代码**。若有讨论，请在下方填写同伴姓名。

**讨论同学**：请填写姓名（如无可写 `None`）

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## Part 1：Jupyter 基础技巧

### 查看函数文档

使用 `help()` 可以查看函数的官方说明。

```python
help(print)
```

你也可以在 Notebook 中直接查看文档。

将光标放在下方 `print` 上，按 `Shift + Tab` 查看函数提示。

```python
print('Welcome to Data 100.')
```

### 导入库与 Magic 命令

课程中会频繁使用数据分析常用库。通常在 Notebook 顶部统一导入，并使用约定俗成的别名。

```python
# DO NOT EDIT THIS CELL
import pandas as pd
import numpy as np
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
%matplotlib inline
```

`%matplotlib inline` 会让 `matplotlib` 图像直接显示在 Notebook 中。

另一个常用命令是 `%%time`，可统计整个单元格执行时间。
说明：`%%` 是单元魔法命令（作用于整格），`%` 是行魔法命令（作用于单行）。

```python
%%time
lst = []
for i in range(100):
    lst.append(i)
```

### 常用快捷键

建议熟悉以下操作：
1. `Shift + Enter`：运行当前单元并跳到下一格
2. `Ctrl + Enter`：仅运行当前单元
3. `Ctrl + /`：批量注释/取消注释
4. `Esc` 进入命令模式后：`a` 上方插入、`b` 下方插入、`dd` 删除单元

### 运行单元格

除快捷键外，也可使用左上角 `Run` 按钮。
当你重启内核后，`Run All Above Selected Cell` 很适合快速恢复到当前进度。

<br><br>

<hr style="border: 1px solid #fdb515;" />

## Part 2：预备知识与练习

请按题目顺序完成。每道题后请自行运行并检查结果。
请勿删除题目与作答单元。完成后建议执行一次“Restart & Run All”。

### Python 复习

默认你已具备基础 Python 语法。若需要复习，可参考：

- [Python Tutorial](https://docs.python.org/3/tutorial/index.html)
- [Composing Programs Chapter 1](https://www.composingprograms.com/pages/11-getting-started.html)

### `NumPy` 复习

`NumPy` 是本课程进行数值计算的核心工具。可快速回顾：

- [NumPy Quick Start Tutorial](https://numpy.org/doc/stable/user/quickstart.html)
- [Data 8 NumPy Chapter](https://www.inferentialthinking.com/chapters/05/1/Arrays)

<br/><br/>

---

### 题目 1

`NumPy` 的核心数据结构是数组（array）。请使用 `np.array` 创建数组 `arr`，其元素按顺序为：`1, 2, 3, 4, 5`。

```python
arr = ...
arr
```

除了访问元素外，也可查看数组属性（如 `shape`、`dtype`）。下方代码用于演示索引与属性读取。

```python
arr[3]
```

```python
arr[2:4]
```

```python
arr.shape
```

```python
arr.dtype
```

与 Python 列表不同，`NumPy` 数组要求元素类型一致；必要时会自动进行类型转换。

```python
# A regular Python list can store items of different data types
[1, '3']
```

```python
# Arrays will convert everything to the same data type
np.array([1, '3'])
```

```python
# Another example of array type conversion
np.array([5, 8.3])
```

数组支持向量化运算：同长度数组之间的算术运算会按元素逐一计算。

```python
# Python list addition will concatenate the two lists
[1, 2, 3] + [4, 5, 6]
```

```python
# NumPy array addition will add them element-wise
np.array([1, 2, 3]) + np.array([4, 5, 6])
```

### 题目 2

<br/><br/>

---

#### 题目 2a

编写函数 `summation(n)`，计算（其中 $n \ge 1$）：

$$\sum_{i=1}^{n} i^3 + 3 i^2$$

要求：不要使用 `for` 循环，可参考 `np.arange`。

```python
def summation(n):
    """Compute the summation i^3 + 3 * i^2 for 1 <= i <= n."""
    ...
```

<br/><br/>

---

### 题目 2b

编写函数 `elementwise_array_sum(list_1, list_2)`：
- 对 `list_1` 每个元素求平方
- 对 `list_2` 每个元素求立方
- 返回两者按元素相加后的结果

要求：输入是 Python 列表，输出是 `NumPy` 数组；不要使用 `for` 循环。

```python
def elementwise_array_sum(list_1, list_2):
    """Compute x^2 + y^3 for each x, y in list_1, list_2. 
    
    Assume list_1 and list_2 have the same length.
    
    Return a NumPy array.
    """
    assert len(list_1) == len(list_2), "both args must have the same number of elements"
    ...
```

下面给出一个不使用 `NumPy` 的实现，用于与向量化写法做性能对比。

```python
def elementwise_list_sum(list_1, list_2):
    """Compute x^2 + y^3 for each x, y in list_1, list_2. 
    
    Assume list_1 and list_2 have the same length.
    """

    return [x ** 2 + y ** 3 for x, y in zip(list_1, list_2)]
```

在样本较小（如 10 个元素）时，两种实现耗时可能接近。

```python
sample_list_1 = list(range(10))
sample_array_1 = np.arange(10)
```

```python
%%time
elementwise_list_sum(sample_list_1, sample_list_1)
```

```python
%%time
elementwise_array_sum(sample_array_1, sample_array_1)
```

当数据规模增大时，向量化计算通常优势明显。

```python
sample_list_2 = list(range(100000))
sample_array_2 = np.arange(100000)
```

```python
%%time
elementwise_list_sum(sample_list_2, sample_list_2)
# The semicolon hides the output
; 
```

```python
%%time
elementwise_array_sum(sample_array_2, sample_array_2)
# The semicolon hides the output
; 
```

结论：在大规模数据处理中，`NumPy` 向量化通常能显著提升运行效率。

<br/><br/>

---

### 题目 2c

已知总体方差公式：

$$\sigma^2 = \frac{\sum_{i=1}^N (x_i - \mu)^2}{N}$$

请补全 `mean(population)` 与 `variance(population)`。
要求：不要使用 `np.var` 与 `np.mean`，且不要使用 `for` 循环。

```python
def mean(population):
    """
    Returns the mean of population (mu)
    
    Keyword arguments:
    population -- a numpy array of numbers
    """
    # Calculate the mean of a population
    ...

def variance(population):
    """
    Returns the variance of population (sigma squared)
    
    Keyword arguments:
    population -- a numpy array of numbers
    """
    # Calculate the variance of a population
    ...
```

<br/><br/>

---

### 题目 2d

已生成随机数组 `random_arr`。请将满足 $2x^4 > 1$ 的元素筛选到 `valid_values`。

要求：不要使用 `for` 循环，建议使用布尔索引（Boolean Indexing）。

```python
np.random.seed(42)
random_arr = np.random.rand(60)
valid_values = ...
valid_values
```

<br><br>

<hr style="border: 1px solid #fdb515;" />

## Part 3：绘图实践

本部分使用 `matplotlib` 与 `NumPy` 完成函数与切线可视化。

### 题目 3

考虑函数 $f(x)=x^2$，其中 $-\infty < x < \infty$。

<br/><br/>

---

### 题目 3a

求 $f$ 在 $x=0$ 处的切线方程，并在 Markdown 单元中写出答案。

<br/><br/>

---

### 题目 3b

求 $f$ 在 $x=8$ 处的切线方程，并在 Markdown 单元中写出答案。

<br/><br/>

---

### 题目 3c

编写代码绘制：$f(x)$、$x=0$ 处切线、$x=8$ 处切线。

绘图要求：
- `x` 轴范围：`(-15, 15)`
- `y` 轴范围：`(-100, 300)`
- 图像尺寸：`(5, 5)`

可使用：`np.linspace`、`plt.plot`、`plt.figure`、`plt.ylim` 等函数。

```python
def f(x):
    ...
    
def f_deriv(x):
    ...

def plot(f, f_deriv):
    # Hint: Use np.linspace()
    x = ...
    ...

plot(f, f_deriv)
```

<br/><br/>

---

### 题目 4（不计分）

阅读并运行下方 3D 绘图代码，将每处 `# Your answer here` 替换为该行代码功能说明（含参数含义）。
建议结合 `Shift + Tab` 查看函数文档。

```python
from mpl_toolkits.mplot3d import axes3d

u = np.linspace(1.5 * np.pi, -1.5 * np.pi, 100)
# Your answer here
[x, y] = np.meshgrid(u, u)
# Your answer here
squared = np.sqrt(x.flatten() ** 2 + y.flatten() ** 2)
z = np.cos(squared)
# Your answer here
z = z.reshape(x.shape)
# Your answer here

fig = plt.figure(figsize = (6, 6))
ax = fig.add_subplot(111, projection = '3d')
# Your answer here
ax.plot_wireframe(x, y, z, rstride = 5, cstride = 5, lw = 2)
# Your answer here
ax.view_init(elev = 60, azim = 25)
# Your answer here
plt.savefig("figure1.png")
# Your answer here
```

<br/><br/>
<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 恭喜完成 Lab 01

### 提交说明

运行下方导出单元后会生成用于提交的压缩文件。若导出失败，请先确认所有必需单元可正常运行。

## 提交

请先按顺序运行全部单元，确认图像/输出完整后，再执行导出单元。

Tests for q1 (points: 1)
```
>>> type(arr) is np.ndarray
True
```
```
>>> all(arr == np.array(np.array([1, 2, 3, 4, 5])))
True
```


Tests for q2a (points: 1)
```
>>> bool(summation(1) == 4)
True
```
```
>>> bool(summation(2) == 24)
True
```
```
>>> bool(summation(100) == 26517550)
True
```


Tests for q2b (points: 1)
```
>>> all(elementwise_array_sum([], []) == np.array([]))
True
```
```
>>> all(elementwise_array_sum([1], [1]) == np.array([2]))
True
```
```
>>> all(elementwise_array_sum([-1], [1]) == np.array([2]))
True
```
```
>>> all(elementwise_array_sum([1], [-1]) == np.array([0]))
True
```
```
>>> all(elementwise_array_sum([1, 2, 3], [1, 2, 3]) == np.array([2, 12, 36]))
True
```
```
>>> all(elementwise_array_sum([1, 5, 2], [3, 6, 6]) == np.array([28, 241, 220]))
True
```
```
>>> type(elementwise_array_sum([], [])) is np.ndarray
True
```


Tests for q2c (points: 1)
```
>>> population_0 = np.random.randn(100)
>>> bool(np.isclose(mean(population_0), np.mean(population_0), atol=1e-06))
True
```
```
>>> population_1 = 3 * np.random.randn(100) + 5
>>> bool(np.isclose(mean(population_1), np.mean(population_1), atol=1e-06))
True
```
```
>>> population_0 = np.random.randn(100)
>>> bool(np.isclose(variance(population_0), np.var(population_0), atol=1e-06))
True
```
```
>>> population_1 = 3 * np.random.randn(100) + 5
>>> bool(np.isclose(variance(population_1), np.var(population_1), atol=1e-06))
True
```


Tests for q2d (points: 1)
```
>>> len(valid_values) == 10
True
```
```
>>> np.allclose(valid_values, [0.95071431, 0.86617615, 0.96990985, 0.94888554, 0.96563203, 0.9093204, 0.96958463, 0.93949894, 0.89482735, 0.92187424])
True
```
